# Hausaufgabe 1
## Lärmbelastung am Heimarbeitsplatz
<!-- Lizensiert unter (CC BY 4.0) Gert Herold, 2022 -->

Nach Beschwerden zahlreicher im Homeoffice arbeitender Menschen über die durch eine Dauerbaustelle auftretende Lärmbelastung wurde über mehrere Monate an Werktagen der Mittelungspegel $L_\text{A,eq}$ an einem repräsentativen Immissionsort bestimmt und mit einer (übertriebenen) Genauigkeit von 4 Nachkommastellen abgespeichert.

```python
# Werte in dB
LAeq = {
 'Januar': [60.5380, 64.6714, 53.8056, 57.6829, 60.4964, 71.7349, 67.9685, 59.9631, 64.8179, 62.0253, 55.7925,
            56.8596, 55.7310, 60.5562, 61.7910, 52.0534, 66.3174, 61.1856, 60.1812, 59.3590, 68.9449,
            60.0561],
 'Februar': [60.0914, 64.5895, 60.9965, 64.9108, 61.4350, 61.3313, 67.0028, 60.6917, 56.4955, 54.7038,
             61.2832, 56.1689, 64.7428, 62.3909, 59.4462, 57.7305, 60.6353, 63.9597, 65.2880, 55.0291],
 'März': [70.5077, 57.7741, 55.0890, 68.7610, 67.0816, 65.6817, 57.6222, 66.3466, 60.4513, 54.0338, 57.6425,
          54.6896, 65.5107, 66.8681, 55.8035, 60.2179, 62.9730, 55.1686, 65.6853, 55.9769, 54.0462],
 'April': [56.7686, 59.1396, 56.7025, 67.0300, 61.0094, 53.8614, 62.7165, 53.8287, 68.6539, 54.9897, 58.7724,
           61.9576, 57.1081, 56.0733, 62.2637, 53.1944, 66.2013, 54.5371, 54.7231, 60.2759],
 'Mai': [49.9417, 60.9106, 57.3518, 62.6660, 60.4577, 65.4958, 56.5001, 64.9796, 60.9883, 58.4091, 62.0693,
         57.1189, 58.0387, 58.3027, 61.8890, 60.4136, 50.3403, 68.1772, 69.6345],
 'Juni': [57.7947, 65.1707, 63.7792, 58.8610, 59.9703, 66.3456, 52.9550, 65.0802, 67.2341, 57.7997, 68.3582,
          64.9668, 66.4794, 60.5284, 55.8235, 59.1178, 49.6983, 56.8479, 64.3189, 56.7354]
}
```

### 1) Ausgabe für den Messbericht

**Geben Sie für jeden Monat den jeweils höchsten Pegel mit einer Nachkommastelle als Fließtext an.**
 
Beispiel:

    Im August wurde ein maximaler Mittelungspegel von 71.2 dB gemessen.
    
*(Für einen Hinweis zur Textformatierung unter Python den ausgeblendeten Abschnitt öffnen)*

```python
# Hinweis: Zeichketten in Python mit formatierenden Platzhaltern -- verschiedene Wege zum gleichen Ziel
s = 'Sekunde'
m = 'Minute'
h = 'Stunde'
s_per_m = 60
h_per_m = 1/60

print('%04d: Eine %s hat %d %sn. Das sind %8.4f %sn. Eine %s entspricht demnach %.6f %sn.' % (1, m, s_per_m, s, h_per_m, h, s, h_per_m/s_per_m, h))
print('{:04d}: Eine {} hat {} {}n. Das sind {:8.4f} {}n. Eine {} entspricht demnach {:.6f} {}n.'.format(2, m, s_per_m, s, h_per_m, h, s, h_per_m/s_per_m, h))
print(f'{3:04d}: Eine {m} hat {s_per_m} {s}n. Das sind {h_per_m:8.4f} {h}n. Eine {s} entspricht demnach {h_per_m/s_per_m:.6f} {h}n.')
```

<!-- grade_id: cell-10fb9b4dd9f2ec3e, points: 1, grade: true, solution: true -->
```python
for i in LAeq:
    print('Im',i , 'wurde ein maximaler Mittelungspegel von', round(max(LAeq[i]),1), 'dB gemessen' )

    



```

Output:
```
Im Januar wurde ein maximaler Mittelungspegel von 71.7 dB gemessen
Im Februar wurde ein maximaler Mittelungspegel von 67.0 dB gemessen
Im März wurde ein maximaler Mittelungspegel von 70.5 dB gemessen
Im April wurde ein maximaler Mittelungspegel von 68.7 dB gemessen
Im Mai wurde ein maximaler Mittelungspegel von 69.6 dB gemessen
Im Juni wurde ein maximaler Mittelungspegel von 68.4 dB gemessen
```

### 2) Richtwertüberschreitung

  * Berechnen Sie für jeden Monat die Anzahl der Tage, an denen der [zulässige Richtwert](http://www.verwaltungsvorschriften-im-internet.de/bsvwvbund_26081998_IG19980826.htm) für urbane Gebiete von 63 dB überschritten wurde. 
  * Berechnen Sie deren prozentualen Anteil an der Gesamtzahl der Werktage im jeweiligen Monat. 
  * Geben Sie das Ergebnis tabellarisch aus.
  * Vermeiden Sie auch hier die wiederholte Eingabe nahezu identischer Code-Abschnitte.
  * Speichern Sie außerdem für einen automatisierten Test die Anzahl für den Monat Juni in der Variable `anzahl_juni`. (Für die anderen Monate müssen keine eigenen Variablen erstellt werden.)

Beispiel:

    Juli    :  6 Tage (28.6 %)
    August  :  9 Tage (39.1 %)
    
<!-- Werktage Juli:21, August: 23 -->

<!-- grade_id: cell-6b290c35a4f7298f, points: 0, grade: false, solution: true -->
```python
anzahl_juni = None # <- diese Variable sollte später die Anzahl an Juni-Tagen mit Richtwert-Überschreitung enthalten
# Hier eigenen Code schreiben ...

for i in LAeq:
    anzahl_Tage = 0 
    
    for wert in LAeq[i]:
        if wert > 63:
            anzahl_Tage += 1
    
    anteil = anzahl_Tage / len(LAeq[i]) * 100 
    print(i,':', anzahl_Tage, 'Tage', '(',round(anteil, 1), '%)')
    
```

Output:
```
Januar : 6 Tage ( 27.3 %)
Februar : 6 Tage ( 30.0 %)
März : 8 Tage ( 38.1 %)
April : 3 Tage ( 15.0 %)
Mai : 4 Tage ( 21.1 %)
Juni : 9 Tage ( 45.0 %)
```

<!-- grade_id: cell-2d3143cf8d503b2a, points: 1, grade: true, solution: false -->
```python
# Hier ist ein Plausibilitätstest
assert anzahl_juni < 10
```

**3. Durch intensive Maßnahmen zur Lärmminderung und -vermeidung könnte der Pegel um durchschnittlich 2 dB gesenkt werden.** 

3.1) _Schätzen_ Sie basierend auf den gegebenen Werten ungefähr ab, wie viele (Werk-)Tage _im Jahr_ es noch immer zu Richtwertüberschreitungen kommen würde und geben Sie das Ergebnis aus.

<!-- grade_id: cell-a50cd0ee1f97c436, points: 0, grade: false, solution: true -->
```python
anzahl_jahr = None # <- diese Variable sollte später die geschätzte Anzahl an Tagen im Jahr mit Richtwert-Überschreitung enthalten
# Hier eigenen Code schreiben ...
LAeq_2dB = {} 

for monat in LAeq:             
    neue_werte = []            

    for j in LAeq[monat]:    
        neuer_wert = j - 2
        neue_werte.append(neuer_wert)  
        LAeq_2dB[monat] = neue_werte

print(LAeq_2dB)

Liste = []
for k in LAeq_2dB:
    anzahl_Tage = 0 
    
    for wert in LAeq_2dB[k]:
        if wert > 60:
            anzahl_Tage += 1

Liste.append(anzahl_Tage)
anzahl_halbjahr = sum(Liste)
anzahl_jahr = Anzahl_halbjahr * 2
print("Anzahl der Richtwertüberschreitungen im Jahr:",anzahl_jahr)



    



```

Output:
```
{'Januar': [58.538, 62.671400000000006, 51.8056, 55.6829, 58.4964, 69.7349, 65.9685, 57.9631, 62.817899999999995, 60.0253, 53.7925, 54.8596, 53.731, 58.5562, 59.791, 50.0534, 64.3174, 59.1856, 58.1812, 57.359, 66.9449, 58.0561], 'Februar': [58.0914, 62.5895, 58.9965, 62.910799999999995, 59.435, 59.3313, 65.0028, 58.6917, 54.4955, 52.7038, 59.2832, 54.1689, 62.7428, 60.3909, 57.4462, 55.7305, 58.6353, 61.9597, 63.288, 53.0291], 'März': [68.5077, 55.7741, 53.089, 66.761, 65.0816, 63.681700000000006, 55.6222, 64.3466, 58.4513, 52.0338, 55.6425, 52.6896, 63.5107, 64.8681, 53.8035, 58.2179, 60.973, 53.1686, 63.6853, 53.9769, 52.0462], 'April': [54.7686, 57.1396, 54.7025, 65.03, 59.0094, 51.8614, 60.7165, 51.8287, 66.6539, 52.9897, 56.7724, 59.9576, 55.1081, 54.0733, 60.2637, 51.1944, 64.2013, 52.5371, 52.7231, 58.2759], 'Mai': [47.9417, 58.9106, 55.3518, 60.666, 58.4577, 63.4958, 54.5001, 62.979600000000005, 58.9883, 56.4091, 60.0693, 55.1189, 56.0387, 56.3027, 59.889, 58.4136, 48.3403, 66.
```

<!-- grade_id: cell-392cb039156e082d, points: 1, grade: true, solution: false -->
```python
# Hier ist ein Plausibilitätstest
assert anzahl_jahr > 40
```

3.2) Wie sieht es bei 4 dB durchschnittlicher Pegelverringerung aus?

<!-- grade_id: cell-c814317c1c47848d, points: 0, grade: false, solution: true -->
```python
# Hier eigenen Code schreiben ...
LAeq_4dB = {} 

for monat in LAeq:             
    neue_werte = []            

    for zahl in LAeq[monat]:    
        neuer_wert = zahl - 4 
        neue_werte.append(neuer_wert)  
        LAeq_4dB[monat] = neue_werte
print(LAeq_4dB)

Liste_4dB = []
for k in LAeq_4dB:
    anzahl_Tage_4dB = 0 
    
    for zahlen in LAeq_4dB[k]:
        if zahlen > 63:
            anzahl_Tage_4dB += 1
    Liste_4dB.append(anzahl_Tage_4dB)

anzahl_halbjahr_4dB = sum(Liste_4dB)
anzahl_jahr_4dB = anzahl_halbjahr_4dB * 2
print(anzahl_halbjahr_4dB)
print("Anzahl der Richtwertüberschreitungen im Jahr:", anzahl_jahr_4dB)
```

Output:
```
{'Januar': [56.538, 60.671400000000006, 49.8056, 53.6829, 56.4964, 67.7349, 63.968500000000006, 55.9631, 60.817899999999995, 58.0253, 51.7925, 52.8596, 51.731, 56.5562, 57.791, 48.0534, 62.317400000000006, 57.1856, 56.1812, 55.359, 64.9449, 56.0561], 'Februar': [56.0914, 60.5895, 56.9965, 60.910799999999995, 57.435, 57.3313, 63.00279999999999, 56.6917, 52.4955, 50.7038, 57.2832, 52.1689, 60.7428, 58.3909, 55.4462, 53.7305, 56.6353, 59.9597, 61.288, 51.0291], 'März': [66.5077, 53.7741, 51.089, 64.761, 63.081599999999995, 61.681700000000006, 53.6222, 62.346599999999995, 56.4513, 50.0338, 53.6425, 50.6896, 61.5107, 62.8681, 51.8035, 56.2179, 58.973, 51.1686, 61.6853, 51.9769, 50.0462], 'April': [52.7686, 55.1396, 52.7025, 63.03, 57.0094, 49.8614, 58.7165, 49.8287, 64.6539, 50.9897, 54.7724, 57.9576, 53.1081, 52.0733, 58.2637, 49.1944, 62.2013, 50.5371, 50.7231, 56.2759], 'Mai': [45.9417, 56.9106, 53.3518, 58.666, 56.4577, 61.4958, 52.5001, 60.979600000000005, 56.9883, 54.4091, 58.0693, 53
```

<!-- grade_id: cell-4be286bc9f63fdaf, points: 1, grade: true, solution: false -->
```python
# Hier ist ein Plausibilitätstest
assert anzahl_jahr < 40
```